// variables5.rs
// Make me compile! Execute the command `rustlings hint variables5` if you want a hint :)

// I AM NOT DONE

fn main() {
<<<<<<< HEAD
    let number = "3"; // don't change this line
    println!("Number {}", number);
    let number = 3;
    println!("Number {}", number);
=======
    let number = "T-H-R-E-E";
    println!("Spell a Number : {}", number);
    number = 3;
    println!("Number plus two is : {}", number + 2);
>>>>>>> 2ec0bdf... fix(variables5) : make shadowing more prominent
}
